<problem display_name="Identity Transformation" linenum="386" tex_filename="./block1/lineartransformationsinRn.tex" showanswer="finished" attempts="1" rerandomize="per_student">
  <text>
    <p>
One very special linear transformation is the transformation [mathjaxinline]T: \mathbb R^ n \rightarrow \mathbb R^ n[/mathjaxinline] which is defined simply by [mathjaxinline]T(v) = v[/mathjaxinline] for all [mathjaxinline]v\in \mathbb R^ n[/mathjaxinline]. This is called the identity transformation on [mathjaxinline]\mathbb R^ n[/mathjaxinline], since it takes every vector to itself. </p>
    <p>
What is the standard matrix for the identity transformation on [mathjaxinline]\mathbb R^3[/mathjaxinline]? </p>
    <script linenum="1" filename="./matrixentry2.tex" type="text/python" system_path="python_lib">
def MatrixEntry(expect, ans):
  	import ast
	import numpy as np 
	ret= {'ok':False}
  	atol = 0.01
  	try:
		list_ans = ast.literal_eval(ans)
		list_expect = ast.literal_eval(expect)
  		matrix_ans = np.matrix(list_ans)
  		matrix_expect = np.matrix(list_expect) 
  		if matrix_ans.shape != matrix_expect.shape:
  			ret['msg'] = 'Wrong shape of matrix'
  		elif np.allclose(matrix_ans, matrix_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
  			ret['msg'] = 'Something is wrong'
	except SyntaxError:
		ret['msg'] = 'Wrong input format'
  	return ret
</script>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[1,0,0],[0,1,0],[0,0,1]]">
        <textline correct_answer="[[1,0,0],[0,1,0],[0,0,1]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 We know that the columns of the standard matrix will be [mathjaxinline]T(e_1),T(e_2),T(e_3)[/mathjaxinline]. When [mathjaxinline]T[/mathjaxinline] is the identity transformation, this means that the columns will be [mathjaxinline]e_1, e_2, e_3[/mathjaxinline]. Hence we get the matrix [mathjaxinline]\left[\begin{array}{ccc} 1 &amp;  0 &amp;  0 \\ 0 &amp;  1 &amp;  0 \\ 0 &amp;  0 &amp;  1 \end{array} \right].[/mathjaxinline]  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
