<problem display_name="Two Row Operations" linenum="93" tex_filename="./block1/linearequations.tex" showanswer="finished" attempts="1" rerandomize="per_student">
  <text>
    <p>
Recall the system of equations from the previous question. </p>
    <table id="a0000000007" cellpadding="7" width="100%" cellspacing="0" class="eqnarray" style="table-layout:auto">
      <tr id="a0000000008">
        <td style="width:40%; border:none">&#xA0;</td>
        <td style="vertical-align:middle;                                    text-align:right; border:none">
	    [mathjaxinline]\displaystyle  10x_1 + 30x_2[/mathjaxinline]
        </td>
        <td style="vertical-align:middle;                                    text-align:center; border:none">
	    [mathjaxinline]\displaystyle  =[/mathjaxinline]
        </td>
        <td style="vertical-align:middle;                                    text-align:left; border:none">
	    [mathjaxinline]\displaystyle  10[/mathjaxinline]
        </td>
        <td style="width:40%; border:none">&#xA0;</td>
        <td style="width:20%; border:none" class="eqnnum"/>
      </tr>
      <tr id="a0000000009">
        <td style="width:40%; border:none">&#xA0;</td>
        <td style="vertical-align:middle;                                    text-align:right; border:none">
	    [mathjaxinline]\displaystyle 2x_1 + 6x_2[/mathjaxinline]
        </td>
        <td style="vertical-align:middle;                                    text-align:center; border:none">
	    [mathjaxinline]\displaystyle  =[/mathjaxinline]
        </td>
        <td style="vertical-align:middle;                                    text-align:left; border:none">
	    [mathjaxinline]\displaystyle  3[/mathjaxinline]
        </td>
        <td style="width:40%; border:none">&#xA0;</td>
        <td style="width:20%; border:none" class="eqnnum"/>
      </tr>
    </table>
    <p>
Do the following row operations to the augmented matrix. First, multiply the first row by [mathjaxinline]\frac{1}{10}[/mathjaxinline]. Then, replace the second row with the second row minus twice the first row. What matrix do you get as a result? </p>
    <script linenum="1" filename="./matrixentry2.tex" type="text/python" system_path="python_lib">
def MatrixEntry(expect, ans):
  	import ast
	import numpy as np 
	ret= {'ok':False}
  	atol = 0.01
  	try:
		list_ans = ast.literal_eval(ans)
		list_expect = ast.literal_eval(expect)
  		matrix_ans = np.matrix(list_ans)
  		matrix_expect = np.matrix(list_expect) 
  		if matrix_ans.shape != matrix_expect.shape:
  			ret['msg'] = 'Wrong shape of matrix'
  		elif np.allclose(matrix_ans, matrix_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
  			ret['msg'] = 'Something is wrong'
	except SyntaxError:
		ret['msg'] = 'Wrong input format'
  	return ret
</script>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[1,3,1],[0,0,1]]">
        <textline correct_answer="[[1,3,1],[0,0,1]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
The result is </p>
          <table id="a0000000010" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{cc:c} 1 &amp;  3 &amp;  1 \\ 0 &amp;  0 &amp;  1 \\ \end{array} \right].[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
        </font>
      </solution>
    </p>
  </text>
</problem>
