<problem display_name="Solving Matrix-Vector Equations" linenum="173" tex_filename="./block1/matrixtimesvector.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
Let </p>
    <table id="a0000000053" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]A = \left[ \begin{array}{cc} 1 &amp;  2 \\ 1 &amp;  3 \\ 3 &amp;  -1 \end{array} \right] ,[/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
 and let [mathjaxinline]v = \left[\begin{array}{c} 3 \\ 2 \\ 16 \end{array} \right].[/mathjaxinline] The equation [mathjaxinline]Ax = v[/mathjaxinline] has exactly one solution &#x2013; what is it? </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[5],[-1]]">
        <textline correct_answer="[[5],[-1]]"/>
      </customresponse>
    </p>
    <p>
Let [mathjaxinline]w = \left[\begin{array}{c} 3 \\ -4 \\ 7 \end{array} \right].[/mathjaxinline] With the same [mathjaxinline]A[/mathjaxinline] as above, how many solutions does [mathjaxinline]Ax = w[/mathjaxinline] have? </p>
    <p>
      <optionresponse>
        <optioninput options="('Exactly one solution','More than one solution','No solutions')" correct="No solutions"/>
      </optionresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 We will row reduce the augmented matrix: <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 1 &amp;  3 &amp;  2\\ 3 &amp;  -1 &amp;  16 \end{array} \right][/mathjaxinline] <br/><br/>Multiply the first row by -1, and add this to the second row: <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 0 &amp;  1 &amp;  -1\\ 3 &amp;  -1 &amp;  16 \end{array} \right][/mathjaxinline] <br/><br/>Multiply the first row by -3, and add this to the third row: <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 0 &amp;  1 &amp;  -1\\ 0 &amp;  -7 &amp;  7 \end{array} \right][/mathjaxinline] <br/><br/>Multiply the second row by 7, and add this to the third row: <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 0 &amp;  1 &amp;  -1\\ 0 &amp;  0 &amp;  0 \end{array} \right][/mathjaxinline] <br/><br/>Multiply the second row by -2, and add this to the first row: <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  0 &amp;  5 \\ 0 &amp;  1 &amp;  -1\\ 0 &amp;  0 &amp;  0 \end{array} \right][/mathjaxinline] <br/><br/>Our vector [mathjaxinline]x[/mathjaxinline] is given by the last column. <br/><br/>We can perform the same steps to answer the second question:<br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 1 &amp;  3 &amp;  -4\\ 3 &amp;  -1 &amp;  7 \end{array} \right][/mathjaxinline] <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 0 &amp;  1 &amp;  -7\\ 3 &amp;  -1 &amp;  7 \end{array} \right][/mathjaxinline] <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 0 &amp;  1 &amp;  -7\\ 0 &amp;  -7 &amp;  -2 \end{array} \right][/mathjaxinline] <br/><br/>[mathjaxinline]\left[\begin{array}{cc:c} 1 &amp;  2 &amp;  3 \\ 0 &amp;  1 &amp;  -7\\ 0 &amp;  0 &amp;  -51 \end{array} \right][/mathjaxinline] <br/><br/>We can stop here. The last row of the augmented matrix implies the equation [mathjaxinline]x_1 \cdot 0 + x_2 \cdot 0 = -51[/mathjaxinline], which has no solution. Therefore, the original equation has no solution.  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
