<problem display_name="List 1" linenum="52" tex_filename="./block1/linearindependence2.tex" showanswer="finished" attempts="5" rerandomize="per_student">
  <text>
    <script linenum="54" filename="./block1/linearindependence2.tex" type="text/python" system_path="python_lib">

def test_ld_1(expect,ans):
    import ast
    import numpy as np 
    if (ans=='No') or (ans=='NO') or (ans=='no'):
        return {'ok': False, 'msg': 'The answer is not no'}
    list_expect = ast.literal_eval('[[1,-1,0,0],[2,-1,3,5],[-3,3,0,0]]')
    list_vec = []
    for vec in list_expect:
        list_vec.append(np.array(vec))
    list_scalars = [float(s) for s in ans.split(',')]
    ret = {'ok':False}
    # check scalars are not all zero
    max_sca = max(list_scalars)
    min_sca = min(list_scalars)
    if (max_sca == 0) and (min_sca ==0):
        ret['msg'] = 'Your answer should contain at least one non-zero scalar.'
        return ret
    # check there are enough scalars
    if (len(list_scalars) != len(list_vec)):
        ret['msg'] = 'Your answer contains too many/too few scalars.'
        return ret    
    # check if a0*v0 + a1*v1 + ... + a_n*vn = 0
    n = len(list_scalars)
    sum = 0
    for i in range(0,n):
        sum += list_scalars[i] * list_vec[i]
    print(sum)
    if np.allclose(sum, np.zeros(4)):
        ret['ok'] = True
    return ret  
</script>
    <p>
Consider the list of vectors [mathjaxinline]\{ v_1; v_2; v_3\}[/mathjaxinline], where </p>
    <table id="a0000000002" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]v_1 = \left[\begin{array}{c} 1 \\ -1 \\ 0 \\ 0 \end{array} \right], v_2 = \left[\begin{array}{c} 2 \\ -1 \\ 3 \\ 5 \end{array} \right], v_3 = \left[\begin{array}{c} -3 \\ 3 \\ 0 \\ 0 \end{array} \right].[/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
Can you find scalars [mathjaxinline]a_1, a_2, a_3[/mathjaxinline], not all zero, such that [mathjaxinline]a_1v_1 + a_2 v_2 + a_3 v_3 = \mathbf{0}[/mathjaxinline]? If so, enter an example of such scalars below. If not, enter &#x2018;No' </p>
    <p>
      <customresponse cfn="test_ld_1" expect="3,0,1">
        <textline correct_answer="3,0,1"/>
      </customresponse>
    </p>
    <p>
Is the list [mathjaxinline]\{ v_1; v_2; v_3\}[/mathjaxinline] linearly dependent or linearly independent? </p>
    <p>
      <optionresponse>
        <optioninput options="('Linearly dependent','Linearly independent')" correct="Linearly dependent"/>
      </optionresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 We notice that [mathjaxinline]-3v_1 = v_3[/mathjaxinline], so that [mathjaxinline]3 v_1 + 0 v_2 + 1 v_3 = \mathbf{0}[/mathjaxinline]. Therefore, [mathjaxinline]a_1 = 3[/mathjaxinline], [mathjaxinline]a_2 = 0[/mathjaxinline], and [mathjaxinline]a_3 = 1[/mathjaxinline] is a possible set of scalars that satisfies the equation. (Any scalar multiple of that triple is also a set that satisfies the equation.) </p>
          <p>
The list [mathjaxinline]\{ v_1; v_2; v_3\}[/mathjaxinline] is linearly dependent because, as we've just seen, there exist sets of coefficients, not all zero, that yield [mathjaxinline]a_1v_1 + a_2 v_2 + a_3 v_3 = \mathbf{0}[/mathjaxinline].  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
