<problem display_name="Compute the Products" linenum="236" tex_filename="./block2/matrixmult.tex" showanswer="finished" attempts="1" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]A = \left[\begin{array}{cc} 1 &amp;  -1 \\ 2 &amp;  0 \end{array} \right][/mathjaxinline] and [mathjaxinline]B = \left[\begin{array}{cc} 0 &amp;  3 \\ 1 &amp;  2 \end{array} \right][/mathjaxinline]. </p>
    <p>
      <div class="hideshowbox">
        <h4 onclick="hideshow(this);" style="margin: 0px">How to Enter a Matrix<span class="icon-caret-down toggleimage"/></h4>
        <div class="hideshowcontent">
          <p>
To enter the matrix [mathjaxinline]\left[ \begin{array}{ccc} 0&amp; 1&amp; 2 \\ 1&amp; 2&amp; 3 \end{array} \right][/mathjaxinline], type [[0,1,2],[1,2,3]]. (Type the same thing if you want to enter the augmented matrix [mathjaxinline]\left[ \begin{array}{cc:c} 0&amp; 1&amp; 2 \\ 1&amp; 2&amp; 3 \end{array} \right][/mathjaxinline].) </p>
          <p>
Use decimals only. </p>
        </div>
        <p class="hideshowbottom" onclick="hideshow(this);" style="margin: 0px">
          <a href="javascript: {return false;}">Show</a>
        </p>
      </div>
    </p>
    <script linenum="1" filename="./matrixentry2.tex" type="text/python" system_path="python_lib">
def MatrixEntry(expect, ans):
  	import ast
	import numpy as np 
	ret= {'ok':False}
  	atol = 0.01
  	try:
		list_ans = ast.literal_eval(ans)
		list_expect = ast.literal_eval(expect)
  		matrix_ans = np.matrix(list_ans)
  		matrix_expect = np.matrix(list_expect) 
  		if matrix_ans.shape != matrix_expect.shape:
  			ret['msg'] = 'Wrong shape of matrix'
  		elif np.allclose(matrix_ans, matrix_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
  			ret['msg'] = 'Something is wrong'
	except SyntaxError:
		ret['msg'] = 'Wrong input format'
  	return ret
</script>
    <p>
What is [mathjaxinline]AB[/mathjaxinline]? </p>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[-1,1],[0,6]]">
        <textline correct_answer="[[-1,1],[0,6]]"/>
      </customresponse>
    </p>
    <p>
What is [mathjaxinline]BA[/mathjaxinline]? </p>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[6,0],[5,-1]]">
        <textline correct_answer="[[6,0],[5,-1]]"/>
      </customresponse>
    </p>
    <p>
Is [mathjaxinline]AB[/mathjaxinline] the same as [mathjaxinline]BA[/mathjaxinline]? </p>
    <p>
      <optionresponse>
        <optioninput options="('Yes','No')" correct="No"/>
      </optionresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <table id="a0000000088" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]AB = \left[\begin{array}{cc} 0+-1 &amp;  3-2 \\ 0+0 &amp;  6+0 \end{array} \right] = \left[\begin{array}{cc} -1 &amp;  1 \\ 0 &amp;  6 \end{array} \right].[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
            <br/>
          </p>
          <table id="a0000000089" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]BA = \left[\begin{array}{cc} 0+6 &amp;  0+0 \\ 1+4 &amp;  -1+0 \end{array} \right] = \left[\begin{array}{cc} 6 &amp;  0 \\ 5 &amp;  -1 \end{array} \right].[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
            <br/>
          </p>
          <p>
We observe that [mathjaxinline]AB\neq BA[/mathjaxinline]. </p>
        </font>
      </solution>
    </p>
  </text>
  <SCRIPT src="/static/latex2edx.js" type="text/javascript"/>
  <LINK href="/static/latex2edx.css" rel="stylesheet" type="text/css"/>
</problem>
